---
title       : Homework - Visualizing Whisky Classification
description : In this homework, we have prepared step-by-step instructions for you on how to prepare plots in Bokeh, a library designed for simple and interactive plotting.  We will demonstrate Bokeh by continuing the analysis of Scotch whiskies.
--- type:NormalExercise lang:python xp:100 skills:2 key:07ea54b341

## Exercise 1

This material will be available soon!

*** =instructions
-
*** =hint
-
*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
```

*** =solution
```{python}
```

*** =sct
```{python}
```

