---
title       : Homework - Social Network Analysis
description : Homophily is a network characteristic.  Homophily occurs when nodes that share an edge share a characteristic more often than nodes that do not share an edge.  In these exercises, we will investigate homophily of several characteristics of individuals connected in social networks in rural India.
--- type:NormalExercise lang:python xp:100 skills:2 key:07ea54b341

## Exercise 1

This material will be available soon!

*** =instructions
-
*** =hint
-
*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
```

*** =solution
```{python}
```

*** =sct
```{python}
```









