---
title       : Homework - Caesar Cipher
description : A cipher is a secret code for a language.  In this homework, we will explore a cipher that is reported by contemporary Greek historians to have been used by Julius Caesar to send secret messages to generals during times of war.
--- type:NormalExercise lang:python xp:100 skills:2 key:07ea54b341

## Exercise 1

This material will be available soon!

*** =instructions
-
*** =hint
-
*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
```

*** =solution
```{python}
```

*** =sct
```{python}
```

