---
title       : Case Study 7 - Movie Analysis
description : Description of the movie analysis.




--- type:NormalExercise lang:python xp:100 skills:2 key:e6c40bde56
## Exercise 1


*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```
