---
title       : Module 1 Extra Exercises
description : Exercises for Homework
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf
  
--- type:NormalExercise lang:python xp:100 skills:1
## Exercise 1

This week, we will create a tic-tac-toe (noughts and crosses) simulator and evaluate basic winning strategies.

*** =instructions
- 


*** =hint
- 

*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
# write your code here!
```

*** =solution
```{python}
```

*** =sct
```{python}
test_function("create_board",
              not_called_msg = "Make sure to call `create_board`!",
              incorrect_msg = "Check your definition of `create_board` again.")
test_object("board",
            undefined_msg = "Did you define `board`?",
            incorrect_msg = "It looks like `board` wasn't defined correctly.")
success_msg("Great work!")
```

