---
title       : Case Study 3 - Practice with Classification
description : In this case study, we will analyze a dataset consisting of an assortment of wines classified into "high quality" and "low quality", and will use k-Nearest Neighbors to predict whether or not other information about the wine helps us correctly guess whether a new wine will be of high quality.
--- type:NormalExercise lang:python xp:100 skills:2 key:07ea54b341
## Exercise 1

This material will be available soon!

*** =instructions
-
*** =hint
-
*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
```

*** =solution
```{python}
```

*** =sct
```{python}
```

