---
title       : Module 1 Extra Exercises
description : Exercises for Homework
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf
  
--- type:NormalExercise lang:python xp:100 skills:1 key:07ea54b341
## Exercise 1

A cipher is a secret code for a language.  For these bonus exercises, we will explore a cipher that is reported by contemporary Greek historians to have been used by Julius Caesar to send secret messages to generals during times of war.

*** =instructions
- 


*** =hint
- 

*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
# write your code here!
```

*** =solution
```{python}
```

*** =sct
```{python}
#test_function("create_board",
#              not_called_msg = "Make sure to call `create_board`!",
#              incorrect_msg = "Check your definition of `create_board` again.")
#test_object("board",
#            undefined_msg = "Did you define `board`?",
#            incorrect_msg = "It looks like `board` wasn't defined correctly.")
success_msg("Great work!")
```

