---
title       : Homework - Bird Migration
description : In this homework, we will continue taking a look at patterns of flight for each of the three birds in our dataset.
--- type:NormalExercise lang:python xp:100 skills:2 key:07ea54b341

## Exercise 1

This material will be available soon!

*** =instructions
-
*** =hint
-
*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}
```

*** =solution
```{python}
```

*** =sct
```{python}
```

