---
title       : Module 5 (Network) Extra Exercises
description : Exercises for Homework
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf
  
--- type:NormalExercise lang:python xp:100 skills:1 key:07ea54b341
## Exercise 1

Description

*** =instructions
-  

*** =hint
- 

*** =pre_exercise_code
```{python}
```

*** =sample_code
```{python}

```

*** =solution
```{python}
```

*** =sct
```{python}
test_function("print",
              not_called_msg = "Make sure to call ``!",
              incorrect_msg = "Check your definition of `` again.")
test_object("",
            undefined_msg = "Did you define ``?",
            incorrect_msg = "It looks like `` wasn't defined correctly.")
success_msg("Great work!")
```
