---
title       : Case Study 7 - Movie Analysis
description : Descriptio of the movie analysis.
--- type:NormalExercise lang:python xp:100 skills:2 key:07ea54b341

## Exercise 1

Description of Part 1.

*** =instructions
-  Do `this`.
-  Do `that`.

*** =hint
-  Did you do `this`?

*** =pre_exercise_code
```{python}
data_filepath = "https://s3.amazonaws.com/assets.datacamp.com/production/course_974/datasets/"
```

*** =sample_code
```{python}
import heyyy
```

*** =solution
```{python}
1+1
```

*** =sct
```{python}
test_object("df",
            undefined_msg = "Did you define `df`?",
            incorrect_msg = "It looks like `df` wasn't defined correctly.")
test_object("df1",
            undefined_msg = "Did you define `df1`?",
            incorrect_msg = "It looks like `df1` wasn't defined correctly.")
test_object("df2",
            undefined_msg = "Did you define `df2`?",
            incorrect_msg = "It looks like `df2` wasn't defined correctly.")   
test_student_typed("df1.head()",
              pattern=False,
              not_typed_msg="Did you call `df1.head()`?")            
success_msg("Great work!")
```

--- type:NormalExercise lang:python xp:100 skills:2 key:95681c66de
